# Introduction to Quantum Mechanics

##  &#x1f3af; Objective¶
To review basic aspects of quantum mechanics.

## &#x1f4dc; Instructions
Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All).

Make sure you fill in any place that says YOUR CODE HERE or "YOUR ANSWER HERE", as well as your name, username (the prefix to your @university.ext e-mail), and student ID number in the cell below

```python
Name = "First M. Last"
email_user_name = "username"
ID_number = 1234567

import numpy as np
```

## &#x1f914; Which of the following phenomena are strongly associated with the particle-nature of light. 
A. Blackbody radiation <br>
B. Compton Scattering  <br>
C. Electron Diffraction <br>
D. Stern-Gerlach Experiment <br>
E. Photoelectric effect 

```python
from ipywidgets import interact, interactive, fixed, interact_manual
import ipywidgets as widgets

tab = widgets.Tab()
tab.children = [widgets.SelectMultiple(options=['A', 'B', 'C', 'D', 'E']), 
            widgets.SelectMultiple(options=['A', 'B', 'C', 'D', 'E'],value=['A','B','E'],disabled=True)]
tab.set_title(0,'question'); tab.set_title(1,'answer')
tab
```

Output:
```
Tab(children=(SelectMultiple(options=('A', 'B', 'C', 'D', 'E'), value=()), SelectMultiple(disabled=True, index…
```

## What is the frequency of light in Hz ($s^{-1}$) of light with wavelength 500 nm?

<!-- grade_id: autoans4, points: 0, grade: false, solution: true -->
```python
# Use this cell however you like, but report the answer as:
# ans4 = [number]
ans4 = None

### BEGIN SOLUTION
# wavelength * frequency = speed of light
ans4 = 2.998e8/500e-9
### END SOLUTION
```

<!-- grade_id: autotest4, points: 0, grade: true, solution: false -->
```python
assert(isinstance(ans4,float))

### BEGIN HIDDEN TESTS
np.testing.assert_almost_equal(ans4,2.998e8/500e-9,3)
print("the frequency of light with wavelength 500 nm is {0:.3e} Hz".format(ans4))
### END HIDDEN TESTS
```

Output:
```
the frequency of light with wavelength 500 nm is 5.996e+14 Hz
```

The key equations for this problem are:
$$\lambda \nu = c$$
so
$$\nu = \frac{c}{\lambda}$$